# Assignment 2

For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```

<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
import pandas as pd

def proportion_of_education():
    # read the csv file
    df = pd.read_csv('assets/NISPUF17.csv')
    
    # count the number of children with each education level
    less_than_high_school = (df["EDUC1"] == 1)
    high_school = (df["EDUC1"] == 2)
    more_than_high_school = (df["EDUC1"] ==3)
    college = (df["EDUC1"] ==4)

    # calculate the proportion of children with each education level
    total = less_than_high_school + high_school + more_than_high_school + college
    proportion_less_than_high_school = less_than_high_school / total
    proportion_high_school = high_school / total
    proportion_more_than_high_school = more_than_high_school / total
    proportion_college = college / total
    
    # calculate the proportion of children with each education level
    total = less_than_high_school + high_school + more_than_high_school + college
    proportion_less_than_high_school = less_than_high_school / total
    proportion_high_school = high_school / total
    proportion_more_than_high_school = more_than_high_school / total
    proportion_college = college / total

    result = {
        "less than high school": less_than_high_school,
        "high school": high_school,
        "more than high school but not college": more_than_high_school,
        "college": college
    }

    return result

proportion_of_education()
#     # your code goes here
#     # YOUR CODE HERE
# #     raise NotImplementedError()
```

Output:
```
~\Anaconda3\lib\site-packages\pandas\core\ops\array_ops.py in arithmetic_op(left, right, op)
    218         # TODO we should handle EAs consistently and move this check before the if/else
    219         # (https://github.com/pandas-dev/pandas/issues/41165)
--> 220         _bool_arith_check(op, left, right)
    221 
    222         res_values = _na_arithmetic_op(left, right, op)

~\Anaconda3\lib\site-packages\pandas\core\ops\array_ops.py in _bool_arith_check(op, a, b)
    518         ):
    519             op_name = op.__name__.strip("_").lstrip("r")
--> 520             raise NotImplementedError(
    521                 f"operator '{op_name}' not implemented for bool dtypes"
    522             )

NotImplementedError: operator 'truediv' not implemented for bool dtypes
```

```python
import pandas as pd
df = pd.read_csv('assets/NISPUF17.csv')
# set the display options to show all the columns
pd.set_option('display.max_columns', None)
df.head(5)
```

Output:
```
Unnamed: 0  SEQNUMC  SEQNUMHH  PDAT    PROVWT_D     RDDWT_D  STRATUM  YEAR  \
0           1   128521     12852     2         NaN  235.916956     1031  2017   
1           2    10741      1074     2         NaN  957.353840     1068  2017   
2           3   220011     22001     2         NaN  189.611299     1050  2017   
3           4    86131      8613     1  675.430817  333.447418     1040  2017   
4           5   227141     22714     1  482.617748  278.768063     1008  2017   

   AGECPOXR  HAD_CPOX  AGEGRP  BF_ENDR06  BF_EXCLR06  BF_FORMR08  BFENDFL06  \
0       NaN         2       1   395.6875      121.75     182.625        NaN   
1       NaN         2       1        NaN         NaN         NaN        NaN   
2       NaN         2       3        NaN         NaN         NaN        NaN   
3       NaN         2       1        NaN         NaN         NaN        NaN   
4       NaN         2       1   547.8750      273.90     888.000        NaN   

   BFFORMFL06  C1R  C5R  CBF_01  CEN_REG 
```

```python
df['EDUC1']
```

Output:
```
0        4
1        3
2        3
3        4
4        1
        ..
28460    3
28461    2
28462    3
28463    2
28464    4
Name: EDUC1, Length: 28465, dtype: int64
```

```python
df['EDUC1'].max()
```

Output:
```
4
```

```python
df[df['EDUC1'] == 1]
```

Output:
```
Unnamed: 0  SEQNUMC  SEQNUMHH  PDAT    PROVWT_D     RDDWT_D  STRATUM  \
4               5   227141     22714     1  482.617748  278.768063     1008   
15             16    41081      4108     2         NaN  685.693874     1036   
16             17   188611     18861     1  345.352030  196.111859     1057   
38             39   215371     21537     2         NaN  364.501746     1011   
78             79   224821     22482     1  477.408498  386.713197     1056   
...           ...      ...       ...   ...         ...         ...      ...   
28416       28417    20541      2054     2         NaN  119.071194     2052   
28422       28423    81921      8192     1  547.185974  274.357467     2052   
28430       28431   237841     23784     1  537.582011  203.666350     2052   
28431       28432   169591     16959     2         NaN  124.857840     2052   
28444       28445   146991     14699     2         NaN  263.520235     2022   

       YEAR  AGECPOXR  HAD_CPOX  AGEGRP  BF_ENDR06  BF_EXC
```

```python
df[df['EDUC1'] == 2]
```

Output:
```
Unnamed: 0  SEQNUMC  SEQNUMHH  PDAT    PROVWT_D      RDDWT_D  STRATUM  \
8               9   161121     16112     1  978.042448   718.616990     1016   
13             14    46531      4653     2         NaN  1116.347858     1025   
22             23   160991     16099     1  289.125148   300.945768     1041   
26             27    31271      3127     1  519.601548   353.968396     1010   
42             43   205551     20555     2         NaN    96.913664     1056   
...           ...      ...       ...   ...         ...          ...      ...   
28457       28458   117081     11708     2         NaN    26.503010     2004   
28458       28459   120601     12060     2         NaN   288.269672     2011   
28459       28460   270341     27034     2         NaN   352.329599     2011   
28461       28462   121401     12140     2         NaN    26.503010     2004   
28463       28464    82921      8292     2         NaN    22.390587     2062   

       YEAR  AGECPOXR  HAD_CPOX  AGEGRP  BF_EN
```

```python
df[df['EDUC1'] == 3]
```

Output:
```
Unnamed: 0  SEQNUMC  SEQNUMHH  PDAT    PROVWT_D      RDDWT_D  STRATUM  \
1               2    10741      1074     2         NaN   957.353840     1068   
2               3   220011     22001     2         NaN   189.611299     1050   
5               6   207941     20794     2         NaN    35.995367     1019   
7               8   182861     18286     1  313.225832   163.213401     1001   
9              10   143401     14340     2         NaN  1536.859533     1068   
...           ...      ...       ...   ...         ...          ...      ...   
28450       28451   230881     23088     2         NaN  1007.220224     2022   
28452       28453    82151      8215     1  335.456813   159.895243     2056   
28456       28457    78891      7889     2         NaN    57.484333     2028   
28460       28461    19141      1914     2         NaN    21.188088     2062   
28462       28463   128831     12883     2         NaN    27.520652     2004   

       YEAR  AGECPOXR  HAD_CPOX  AGEGRP  BF_EN
```

```python
df[df['EDUC1'] == 4]
```

Output:
```
Unnamed: 0  SEQNUMC  SEQNUMHH  PDAT    PROVWT_D     RDDWT_D  STRATUM  \
0               1   128521     12852     2         NaN  235.916956     1031   
3               4    86131      8613     1  675.430817  333.447418     1040   
6               7   221381     22138     1  516.427723  248.916038     1008   
11             12    92921      9292     2         NaN   76.306250     1005   
12             13   137041     13704     1   82.582116   45.544717     1051   
...           ...      ...       ...   ...         ...         ...      ...   
28448       28449   237881     23788     2         NaN  390.466235     2076   
28449       28450   132081     13208     1  166.681363   96.330007     2059   
28451       28452   152371     15237     2         NaN   73.286155     2049   
28455       28456   154501     15450     2         NaN   56.655283     2062   
28464       28465   244621     24462     2         NaN   27.520652     2004   

       YEAR  AGECPOXR  HAD_CPOX  AGEGRP  BF_ENDR06  BF_EXC
```

```python
less_than_high_school = (df["EDUC1"] == 1).sum()
high_school = (df["EDUC1"] == 2).sum()
more_than_high_school = (df["EDUC1"] ==3).sum()
college = (df["EDUC1"] ==4).sum()

print(less_than_high_school)
print(high_school)
print(more_than_high_school)
print(college)
```

Output:
```
2904
4906
6999
13656
```

```python
less_than_high_school = (df["EDUC1"] == 1).astype(int).sum()
less_than_high_school
print(type(less_than_high_school))
```

Output:
```
<class 'numpy.int64'>
```

```python
high_school = (df["EDUC1"] == 2).astype(int).sum()
high_school 
```

Output:
```
4906
```

```python
more_than_high_school = (df["EDUC1"] ==3).astype(int).sum()
more_than_high_school
```

Output:
```
6999
```

```python
college = (df["EDUC1"] ==4).astype(int).sum()
college
```

Output:
```
13656
```

```python
df = pd.read_csv('assets/NISPUF17.csv')

# count the number of children with each education level
less_than_high_school = (df["EDUC1"] == 1)
high_school = (df["EDUC1"] == 2)
more_than_high_school = (df["EDUC1"] ==3)
college = (df["EDUC1"] ==4)

# calculate the proportion of children with each education level
total = less_than_high_school + high_school + more_than_high_school + college
proportion_less_than_high_school = less_than_high_school / total
proportion_high_school = high_school / total
proportion_more_than_high_school = more_than_high_school / total
proportion_college = college / total
    
# calculate the proportion of children with each education level
total = less_than_high_school + high_school + more_than_high_school + college
proportion_less_than_high_school = less_than_high_school / total
proportion_high_school = high_school / total
proportion_more_than_high_school = more_than_high_school / total
proportion_college = college / total

result = {
    "less than high school": less_than_high_school,
    "high school": high_school,
    "more than high school but not college": more_than_high_school,
    "college": college
}

```

Output:
```
~\Anaconda3\lib\site-packages\pandas\core\ops\array_ops.py in arithmetic_op(left, right, op)
    218         # TODO we should handle EAs consistently and move this check before the if/else
    219         # (https://github.com/pandas-dev/pandas/issues/41165)
--> 220         _bool_arith_check(op, left, right)
    221 
    222         res_values = _na_arithmetic_op(left, right, op)

~\Anaconda3\lib\site-packages\pandas\core\ops\array_ops.py in _bool_arith_check(op, a, b)
    518         ):
    519             op_name = op.__name__.strip("_").lstrip("r")
--> 520             raise NotImplementedError(
    521                 f"operator '{op_name}' not implemented for bool dtypes"
    522             )

NotImplementedError: operator 'truediv' not implemented for bool dtypes
```

```python
df = pd.read_csv('assets/NISPUF17.csv')

# count the number of children with each education level
less_than_high_school = (df["EDUC1"] == 1).astype(int).sum()
high_school = (df["EDUC1"] == 2).astype(int).sum()
more_than_high_school = (df["EDUC1"] ==3).astype(int).sum()
college = (df["EDUC1"] ==4).astype(int).sum()

# calculate the proportion of children with each education level
total = less_than_high_school + high_school + more_than_high_school + college
proportion_less_than_high_school = less_than_high_school / total
proportion_high_school = high_school / total
proportion_more_than_high_school = more_than_high_school / total
proportion_college = college / total
    
# calculate the proportion of children with each education level
total = less_than_high_school + high_school + more_than_high_school + college
proportion_less_than_high_school = less_than_high_school / total
proportion_high_school = high_school / total
proportion_more_than_high_school = more_than_high_school / total
proportion_college = college / total

result = {
    "less than high school": proportion_less_than_high_school,
    "high school": proportion_high_school,
    "more than high school but not college": proportion_more_than_high_school,
    "college": proportion_college,
}
print(result)
```

Output:
```
{'less than high school': 0.10202002459160373, 'high school': 0.172352011241876, 'more than high school but not college': 0.24588090637625154, 'college': 0.47974705779026877}
```

```python
df = pd.read_csv('assets/NISPUF17.csv')

# count the number of children with each education level
less_than_high_school = (df["EDUC1"] == 1).astype(int).sum()
high_school = (df["EDUC1"] == 2).astype(int).sum()
more_than_high_school = (df["EDUC1"] ==3).astype(int).sum()
college = (df["EDUC1"] ==4).astype(int).sum()
    
# calculate the proportion of children with each education level
total = less_than_high_school + high_school + more_than_high_school + college
proportion_less_than_high_school = less_than_high_school / total
proportion_high_school = high_school / total
proportion_more_than_high_school = more_than_high_school / total
proportion_college = college / total

# Round calculatons
lessThanHS = "{:.2f}".format(proportion_less_than_high_school)
HS = "{:.2f}".format(proportion_high_school)
moreThanHS = "{:.2f}".format(proportion_more_than_high_school)
collegiate = "{:.2f}".format(proportion_college)

result = {
    "less than high school": lessThanHS,
    "high school": HS,
    "more than high school but not college": moreThanHS,
    "college": collegiate,
}

print(result)
```

Output:
```
{'less than high school': '0.10', 'high school': '0.17', 'more than high school but not college': '0.25', 'college': '0.48'}
```

```python
import pandas as pd

def proportion_of_education():
    df = pd.read_csv('assets/NISPUF17.csv')

    # count the number of children with each education level
    less_than_high_school = (df["EDUC1"] == 1).astype(int).sum()
    high_school = (df["EDUC1"] == 2).astype(int).sum()
    more_than_high_school = (df["EDUC1"] ==3).astype(int).sum()
    college = (df["EDUC1"] ==4).astype(int).sum()

    # calculate the proportion of children with each education level
    total = less_than_high_school + high_school + more_than_high_school + college
    proportion_less_than_high_school = less_than_high_school / total
    proportion_high_school = high_school / total
    proportion_more_than_high_school = more_than_high_school / total
    proportion_college = college / total

    # Round calculatons
    less_than_hs = "{:.2f}".format(proportion_less_than_high_school)
    h_s = "{:.2f}".format(proportion_high_school)
    more_than_hs = "{:.2f}".format(proportion_more_than_high_school)
    collegiate = "{:.2f}".format(proportion_college)

    result = {
    "less than high school": less_than_hs,
    "high school": h_s,
    "more than high school but not college": more_than_hs,
    "college": collegiate,
    }
    
    return result

proportion_of_education()
```

Output:
```
{'less than high school': '0.10',
 'high school': '0.17',
 'more than high school but not college': '0.25',
 'college': '0.48'}
```

```python
def get_count(x):
    df = pd.read_csv('assets/NISPUF17.csv')
    df = df["EDUC1"]
    total = df.shape[0]
    mask = df == x
    df.where(mask, inplace = True)
    df = df.dropna()
    return df.shape[0] / total

def proportion_of_education():
    less_high = get_count(1)
    high = get_count(2)
    more_high  = get_count(3)
    college    = get_count(4)
    prop_edu_moms = {"less than high school" :less_high, 
                     "high school" : high, 
                     "more than high school but not college" : more_high, 
                     "college" : college}
    return prop_edu_moms


proportion_of_education()

```

Output:
```
{'less than high school': 0.10202002459160373,
 'high school': 0.172352011241876,
 'more than high school but not college': 0.24588090637625154,
 'college': 0.47974705779026877}
```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

```python
df["DFLU1"]
```

Output:
```
0          NaN
1          NaN
2          NaN
3        192.0
4          NaN
         ...  
28460      NaN
28461      NaN
28462      NaN
28463      NaN
28464      NaN
Name: DFLU1, Length: 28465, dtype: float64
```

```python
df["DFLU2"]
```

Output:
```
0          NaN
1          NaN
2          NaN
3        223.0
4          NaN
         ...  
28460      NaN
28461      NaN
28462      NaN
28463      NaN
28464      NaN
Name: DFLU2, Length: 28465, dtype: float64
```

```python
df["DFLU3"]
```

Output:
```
0          NaN
1          NaN
2          NaN
3        368.0
4          NaN
         ...  
28460      NaN
28461      NaN
28462      NaN
28463      NaN
28464      NaN
Name: DFLU3, Length: 28465, dtype: float64
```

```python
df["DFLU4"].head(20)
```

Output:
```
0    NaN
1    NaN
2    NaN
3    NaN
4    NaN
5    NaN
6    NaN
7    NaN
8    NaN
9    NaN
10   NaN
11   NaN
12   NaN
13   NaN
14   NaN
15   NaN
16   NaN
17   NaN
18   NaN
19   NaN
Name: DFLU4, dtype: float64
```

```python
df["DFLU5"].head(20)
```

Output:
```
0    NaN
1    NaN
2    NaN
3    NaN
4    NaN
5    NaN
6    NaN
7    NaN
8    NaN
9    NaN
10   NaN
11   NaN
12   NaN
13   NaN
14   NaN
15   NaN
16   NaN
17   NaN
18   NaN
19   NaN
Name: DFLU5, dtype: float64
```

```python
df[df["DFLU1"] != 'NaN']
```

Output:
```
Unnamed: 0  SEQNUMC  SEQNUMHH  PDAT    PROVWT_D     RDDWT_D  STRATUM  \
0               1   128521     12852     2         NaN  235.916956     1031   
1               2    10741      1074     2         NaN  957.353840     1068   
2               3   220011     22001     2         NaN  189.611299     1050   
3               4    86131      8613     1  675.430817  333.447418     1040   
4               5   227141     22714     1  482.617748  278.768063     1008   
...           ...      ...       ...   ...         ...         ...      ...   
28460       28461    19141      1914     2         NaN   21.188088     2062   
28461       28462   121401     12140     2         NaN   26.503010     2004   
28462       28463   128831     12883     2         NaN   27.520652     2004   
28463       28464    82921      8292     2         NaN   22.390587     2062   
28464       28465   244621     24462     2         NaN   27.520652     2004   

       YEAR  AGECPOXR  HAD_CPOX  AGEGRP  BF_ENDR06  BF_EXC
```

```python
# Create a Boolean mask for the column of interest
mask = df['DFLU1'].notna()

# Use the mask to filter the DataFrame
filtered_df = df[mask]

filtered_df
```

Output:
```
Unnamed: 0  SEQNUMC  SEQNUMHH  PDAT    PROVWT_D     RDDWT_D  STRATUM  \
3               4    86131      8613     1  675.430817  333.447418     1040   
10             11   128211     12821     1  335.398714  198.123294     1057   
12             13   137041     13704     1   82.582116   45.544717     1051   
16             17   188611     18861     1  345.352030  196.111859     1057   
19             20   136171     13617     1  223.451930  136.912758     1059   
...           ...      ...       ...   ...         ...         ...      ...   
28446       28447   193801     19380     1  321.058203  110.009181     2047   
28447       28448   136841     13684     1  593.426753  337.255316     2031   
28449       28450   132081     13208     1  166.681363   96.330007     2059   
28452       28453    82151      8215     1  335.456813  159.895243     2056   
28454       28455   191451     19145     1   59.626969   29.187316     2074   

       YEAR  AGECPOXR  HAD_CPOX  AGEGRP  BF_ENDR06  BF_EXC
```

```python
filtered_df['DFLU1'].sum()
```

Output:
```
3616794.0
```

```python
(df["CBF_01"]==1).sum()
```

Output:
```
24261
```

```python
(df["CBF_01"]==2).sum()
```

Output:
```
4085
```

```python
(filtered_df["CBF_01"]==1).sum()
```

Output:
```
9930
```

```python
(filtered_df["CBF_01"]==2).sum()
```

Output:
```
1360
```

```python
(filtered_df["CBF_01"]==1).sum() + (filtered_df["CBF_01"]==2).sum()
```

Output:
```
11290
```

```python
(filtered_df["CBF_01"]==1).mean()
```

Output:
```
0.877208480565371
```

```python
(filtered_df["CBF_01"]==2).mean()
```

Output:
```
0.12014134275618374
```

```python
(filtered_df["CBF_01"]==1).mean() + (filtered_df["CBF_01"]==2).mean()
```

Output:
```
0.9973498233215548
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
def average_influenza_doses():
    # Create a Boolean mask for the column of interest
    mask = df['DFLU1'].notna()

    # Use the mask to filter the DataFrame
    filtered_df = df[mask]
    
    # Calculate average number of flu vaccines for each group
    avg_yes_breastmilk = (filtered_df["CBF_01"]==1).mean()
    avg_no_breastmilk = (filtered_df["CBF_01"]==2).mean()
    
    # Round calculatons
    yes = "{:.2f}".format(avg_yes_breastmilk)
    no = "{:.2f}".format(avg_no_breastmilk)
    
    # Return results as a tuple
    new_result = (yes, no)
    
    return new_result

average_influenza_doses()

#     YOUR CODE HERE
#     raise NotImplementedError()
```

Output:
```
('0.88', '0.12')
```

```python
df_bf_inf = df[["CBF_01", "P_NUMFLU"]]

df_bf_inf = df_bf_inf[df_bf_inf["P_NUMFLU"].notna()]

mask_bf = df["CBF_01"] == 1
df_bf = df_bf_inf.where(mask_bf, inplace = False)
df_bf = df_bf[df_bf["CBF_01"].notna()]

mask_nbf = df["CBF_01"] == 2
df_nbf = df_bf_inf.where(mask_nbf, inplace = False)
df_nbf = df_nbf[df_nbf["CBF_01"].notna()]

result = (df_bf["P_NUMFLU"].mean(), df_nbf["P_NUMFLU"].mean())




df_bf_inf.where(mask_bf, inplace = True)
df_bf_inf = df_bf_inf[df_bf_inf["CBF_01"].notna()]
total = df_bf_inf.shape[0]
sum_bf = df_bf_inf["P_NUMFLU"].sum()
avg_bf = sum_bf/total 

def average_influenza_doses():
    df_bf_inf = df[["CBF_01", "P_NUMFLU"]]

    df_bf_inf = df_bf_inf[df_bf_inf["P_NUMFLU"].notna()]

    mask_bf = df["CBF_01"] == 1
    df_bf = df_bf_inf.where(mask_bf, inplace = False)
    df_bf = df_bf[df_bf["CBF_01"].notna()]

    mask_nbf = df["CBF_01"] == 2
    df_nbf = df_bf_inf.where(mask_nbf, inplace = False)
    df_nbf = df_nbf[df_nbf["CBF_01"].notna()]

    result = (df_bf["P_NUMFLU"].mean(), df_nbf["P_NUMFLU"].mean())
    return result

average_influenza_doses()
```

Output:
```
(1.8799187420058687, 1.5963945918878317)
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

```python
df["SEX"]
# 1 == male, 2== female
```

Output:
```
0        1
1        1
2        2
3        2
4        2
        ..
28460    2
28461    2
28462    2
28463    2
28464    2
Name: SEX, Length: 28465, dtype: int64
```

```python
(df["SEX"]==1).sum()
```

Output:
```
14717
```

```python
(df["SEX"]==2).sum()
```

Output:
```
13748
```

```python
df["DVRC1"]
```

Output:
```
0          NaN
1          NaN
2          NaN
3        368.0
4          NaN
         ...  
28460      NaN
28461      NaN
28462      NaN
28463      NaN
28464      NaN
Name: DVRC1, Length: 28465, dtype: float64
```

```python
# Create a Boolean mask for the column of interest
mask2 = df['DVRC1'].notna()

# Use the mask to filter the DataFrame
filtered_df2 = df[mask2]

filtered_df2
```

Output:
```
Unnamed: 0  SEQNUMC  SEQNUMHH  PDAT    PROVWT_D     RDDWT_D  STRATUM  \
3               4    86131      8613     1  675.430817  333.447418     1040   
6               7   221381     22138     1  516.427723  248.916038     1008   
10             11   128211     12821     1  335.398714  198.123294     1057   
12             13   137041     13704     1   82.582116   45.544717     1051   
16             17   188611     18861     1  345.352030  196.111859     1057   
...           ...      ...       ...   ...         ...         ...      ...   
28446       28447   193801     19380     1  321.058203  110.009181     2047   
28447       28448   136841     13684     1  593.426753  337.255316     2031   
28449       28450   132081     13208     1  166.681363   96.330007     2059   
28452       28453    82151      8215     1  335.456813  159.895243     2056   
28454       28455   191451     19145     1   59.626969   29.187316     2074   

       YEAR  AGECPOXR  HAD_CPOX  AGEGRP  BF_ENDR06  BF_EXC
```

```python
(filtered_df2["SEX"]==1).sum()
```

Output:
```
7158
```

```python
(filtered_df2["SEX"]==1)
```

Output:
```
3        False
6        False
10        True
12        True
16        True
         ...  
28446    False
28447     True
28449     True
28452    False
28454    False
Name: SEX, Length: 14076, dtype: bool
```

```python
(filtered_df2["SEX"]==2).sum()
```

Output:
```
6918
```

```python
(filtered_df2["SEX"]==2)
```

Output:
```
3         True
6         True
10       False
12       False
16       False
         ...  
28446     True
28447    False
28449    False
28452     True
28454     True
Name: SEX, Length: 14076, dtype: bool
```

```python
filtered_df2["HAD_CPOX"]
```

Output:
```
3        2
6        2
10       2
12       2
16       2
        ..
28446    2
28447    2
28449    2
28452    2
28454    2
Name: HAD_CPOX, Length: 14076, dtype: int64
```

```python
filtered_df2[filtered_df2["HAD_CPOX"]==1]
```

Output:
```
Unnamed: 0  SEQNUMC  SEQNUMHH  PDAT     PROVWT_D     RDDWT_D  STRATUM  \
461           462   255961     25596     1   523.277248  275.034692     1002   
687           688    72351      7235     1  1849.038009  466.803981     2011   
721           722   176421     17642     1  1643.816096  665.365790     2011   
863           864    39451      3945     1   150.442176   64.267734     2108   
1324         1325    10511      1051     1    54.422973   27.051018     2004   
...           ...      ...       ...   ...          ...         ...      ...   
27720       27721    59601      5960     1   623.665997  300.231789     2041   
28106       28107    53322      5332     1    30.347802   25.141769     2058   
28107       28108    53321      5332     1    36.803372   27.792652     2058   
28314       28315   233711     23371     1    22.125887   13.888073     2007   
28404       28405   177661     17766     1   128.133719  105.089644     2020   

       YEAR  AGECPOXR  HAD_CPOX  AGEGRP  BF_EN
```

```python
filtered_df2[filtered_df2["HAD_CPOX"]==2]
```

Output:
```
Unnamed: 0  SEQNUMC  SEQNUMHH  PDAT    PROVWT_D     RDDWT_D  STRATUM  \
3               4    86131      8613     1  675.430817  333.447418     1040   
6               7   221381     22138     1  516.427723  248.916038     1008   
10             11   128211     12821     1  335.398714  198.123294     1057   
12             13   137041     13704     1   82.582116   45.544717     1051   
16             17   188611     18861     1  345.352030  196.111859     1057   
...           ...      ...       ...   ...         ...         ...      ...   
28446       28447   193801     19380     1  321.058203  110.009181     2047   
28447       28448   136841     13684     1  593.426753  337.255316     2031   
28449       28450   132081     13208     1  166.681363   96.330007     2059   
28452       28453    82151      8215     1  335.456813  159.895243     2056   
28454       28455   191451     19145     1   59.626969   29.187316     2074   

       YEAR  AGECPOXR  HAD_CPOX  AGEGRP  BF_ENDR06  BF_EXC
```

```python
(filtered_df2["HAD_CPOX"]==2) & (filtered_df2["SEX"]==1)
```

Output:
```
3        False
6        False
10        True
12        True
16        True
         ...  
28446    False
28447     True
28449     True
28452    False
28454    False
Length: 14076, dtype: bool
```

```python
vacc_with_cp = filtered_df2[filtered_df2["HAD_CPOX"]==1]
vacc_wo_cp = filtered_df2[filtered_df2["HAD_CPOX"]==2]
```

```python
vacc_with_cp_count = int(vacc_with_cp.shape[0])
vacc_wo_cp_count = int(vacc_wo_cp.shape[0])
print(vacc_with_cp_count)
print(vacc_wo_cp_count)
```

Output:
```
123
13909
```

```python
vacc_total = vacc_with_cp_count + vacc_wo_cp_count
print(vacc_total)
print(type(vacc_total))
```

Output:
```
14032
<class 'int'>
```

```python
vacc_with_cp[vacc_with_cp["SEX"]==1]
```

Output:
```
Unnamed: 0  SEQNUMC  SEQNUMHH  PDAT     PROVWT_D     RDDWT_D  STRATUM  \
461           462   255961     25596     1   523.277248  275.034692     1002   
687           688    72351      7235     1  1849.038009  466.803981     2011   
863           864    39451      3945     1   150.442176   64.267734     2108   
1479         1480   126621     12662     1  1438.229156  722.644266     2010   
1906         1907   223461     22346     1   117.899023   70.800480     2059   
...           ...      ...       ...   ...          ...         ...      ...   
27045       27046    14731      1473     1   134.723345   59.539452     2030   
27720       27721    59601      5960     1   623.665997  300.231789     2041   
28106       28107    53322      5332     1    30.347802   25.141769     2058   
28107       28108    53321      5332     1    36.803372   27.792652     2058   
28404       28405   177661     17766     1   128.133719  105.089644     2020   

       YEAR  AGECPOXR  HAD_CPOX  AGEGRP  BF_EN
```

```python
vacc_with_cp[vacc_with_cp["SEX"]==2]
```

Output:
```
Unnamed: 0  SEQNUMC  SEQNUMHH  PDAT     PROVWT_D      RDDWT_D  STRATUM  \
721           722   176421     17642     1  1643.816096   665.365790     2011   
1324         1325    10511      1051     1    54.422973    27.051018     2004   
1571         1572   148871     14887     1   905.442543   436.332245     2051   
2124         2125   256411     25641     1   599.877193   222.570309     2001   
2627         2628   223241     22324     1   338.739540   181.296459     2060   
3106         3107   107911     10791     1   249.035148   139.068518     2059   
3206         3207    29451      2945     1   145.550120    70.329966     2055   
3572         3573   230231     23023     1   540.851103   292.764148     2054   
3819         3820   120591     12059     1    35.424818    20.988940     2060   
4204         4205     7011       701     1    47.447022    30.549587     2004   
4469         4470   271512     27151     1    40.720180    29.597203     2049   
4560         4561   209101     2091
```

```python
male_vacc_with_cp = vacc_with_cp[vacc_with_cp["SEX"]==1]
female_vacc_with_cp = vacc_with_cp[vacc_with_cp["SEX"]==2]
```

```python
male_vacc_with_cp_count = int(male_vacc_with_cp.shape[0])
female_vacc_with_cp_count = int(female_vacc_with_cp.shape[0])
print(male_vacc_with_cp_count)
print(female_vacc_with_cp_count)
```

Output:
```
68
55
```

```python
vacc_males = filtered_df2[filtered_df2["SEX"]==1]
vacc_males_count = int(vacc_males.shape[0])

vacc_females = filtered_df2[filtered_df2["SEX"]==2]
vacc_females_count = int(vacc_females.shape[0])

print(vacc_males_count)
print(vacc_females_count)
```

Output:
```
7158
6918
```

```python
(vacc_males_count + vacc_females_count) == int(filtered_df2.shape[0])
```

Output:
```
True
```

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
def chickenpox_by_sex():
    
    # Create a Boolean mask for those who received at least 1 varicella dose
    mask2 = df['DVRC1'].notna()

    # Use the mask to filter the DataFrame
    filtered_df2 = df[mask2]
    
    # Isolate those who contracted chickenpox and those who did not into separate DataFrames
    vacc_with_cp = filtered_df2[filtered_df2["HAD_CPOX"]==1]
    vacc_wo_cp = filtered_df2[filtered_df2["HAD_CPOX"]==2]
    
    # Isolate vaccinated males from females into separate DataFrames and generate int variable for each DataFrame
    
    # Males
    vacc_males = filtered_df2[filtered_df2["SEX"]==1]
    vacc_males_count = int(vacc_males.shape[0])
    
    # Females
    vacc_females = filtered_df2[filtered_df2["SEX"]==2]
    vacc_females_count = int(vacc_females.shape[0])
    
    # Isolate males from females who contracted chickenpox into separate DataFrames and generate int variable for each DataFrame
    
    # Males
    male_vacc_with_cp = vacc_with_cp[vacc_with_cp["SEX"]==1]
    male_vacc_with_cp_count = int(male_vacc_with_cp.shape[0])
    
    # Females
    female_vacc_with_cp = vacc_with_cp[vacc_with_cp["SEX"]==2]
    female_vacc_with_cp_count = int(female_vacc_with_cp.shape[0])
    
    # Divide infected count by vaccinated count into separate variables by sex
    male_ratio = male_vacc_with_cp_count / vacc_males_count
    female_ratio = female_vacc_with_cp_count / vacc_females_count 
    
    # Display results as dictionary
    
    cp_results = {
        "male": male_ratio,
        "female": female_ratio
    }
    
    return cp_results

chickenpox_by_sex()
    
#     YOUR CODE HERE
#     raise NotImplementedError()
```

Output:
```
{'male': 0.009499860296172115, 'female': 0.007950274645851402}
```

```python
def chickenpox_by_sex():

    df_cp_sex = df[["SEX","P_NUMVRC","HAD_CPOX"]]
    mask_vax = df_cp_sex["P_NUMVRC"] > 0
    df_cp_sex.where(mask_vax, inplace = True)
    df_cp_sex = df_cp_sex[df_cp_sex["P_NUMVRC"].notna()] 
    df_cp_sex.drop("P_NUMVRC", axis = 1, inplace = True);

    mask_cp = df_cp_sex["HAD_CPOX"]== 1
    df_cp = df_cp_sex.where(mask_cp, inplace = False)
    df_cp = df_cp[df_cp["HAD_CPOX"].notna()]

    mask_ncp = df_cp_sex["HAD_CPOX"]== 2
    df_ncp = df_cp_sex.where(mask_ncp, inplace = False)
    df_ncp = df_ncp[df_ncp["HAD_CPOX"].notna()]

    ar_cp = df_cp.to_numpy()
    ar_ncp = df_ncp.to_numpy()

    # male calculation

    cnt_male_cp = 0
    for i in range (0, ar_cp.shape[0]):
        if(ar_cp[i, 0] == 1):
            cnt_male_cp = cnt_male_cp + 1

    cnt_female_cp = ar_cp.shape[0] - cnt_male_cp

    cnt_male_ncp = 0
    for i in range (0, ar_ncp.shape[0]):
        if(ar_ncp[i, 0] == 1):
            cnt_male_ncp = cnt_male_ncp + 1

    cnt_female_ncp = ar_ncp.shape[0] - cnt_male_ncp

    new_new_result = {"male" : cnt_male_cp/cnt_male_ncp,
              "female" : cnt_female_cp/cnt_female_ncp}
    
    return new_new_result

chickenpox_by_sex()
```

Output:
```
C:\Users\E073016\AppData\Local\Temp\1\ipykernel_11372\4213832900.py:5: SettingWithCopyWarning: 
A value is trying to be set on a copy of a slice from a DataFrame

See the caveats in the documentation: https://pandas.pydata.org/pandas-docs/stable/user_guide/indexing.html#returning-a-view-versus-a-copy
  df_cp_sex.where(mask_vax, inplace = True)
{'male': 0.009675583380762664, 'female': 0.0077918259335489565}
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. The `had_chickenpox_column` is either `1` (for yes) or `2` (for no), and the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine. A positive correlation (e.g., `corr > 0`) means that an increase in `had_chickenpox_column` (which means more nos) would also increase the values of `num_chickenpox_vaccine_column` (which means more doses of vaccine). If there is a negative correlation (e.g., `corr < 0`), it indicates that having had chickenpox is related to an increase in the number of vaccine doses.

Also, `pval` is the probability that we observe a correlation between `had_chickenpox_column` and `num_chickenpox_vaccine_column` which is greater than or equal to a particular value occurred by chance. A small `pval` means that the observed correlation is highly unlikely to occur by chance. In this case, `pval` should be very small (will end in `e-18` indicating a very small number).

[1] This isnt really the full picture, since we are not looking at when the dose was given. Its possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    # this is just an example dataframe
    df=pd.DataFrame({"had_chickenpox_column":np.random.randint(1,3,size=(100)),
                   "num_chickenpox_vaccine_column":np.random.randint(0,6,size=(100))})

    # here is some stub code to actually run the correlation
    corr, pval=stats.pearsonr(df["had_chickenpox_column"],df["num_chickenpox_vaccine_column"])
    
    # just return the correlation
    return corr

corr_chickenpox()

    # YOUR CODE HERE
#     raise NotImplementedError()
```

Output:
```
0.021997266414506722
```

```python
def corr_chickenpox():

    import scipy.stats as stats
    import numpy as np
    import pandas as pd

    df_corr =df[["P_NUMVRC","HAD_CPOX"]]

            # this is just an example dataframe
    df_corr=pd.DataFrame({"had_chickenpox_column":df["HAD_CPOX"],
                           "num_chickenpox_vaccine_column":df["P_NUMVRC"]})

    mask_cp = (df_corr["had_chickenpox_column"] == 1) |(df_corr["had_chickenpox_column"] == 2)
    df_corr.where(mask_cp, inplace = True)
    df_corr = df_corr[df_corr["had_chickenpox_column"].notna()]
    df_corr = df_corr[df_corr["num_chickenpox_vaccine_column"].notna()]



            # here is some stub code to actually run the correlation
    corr, pval=stats.pearsonr(df_corr["had_chickenpox_column"],df_corr["num_chickenpox_vaccine_column"])
    
    return corr

corr_chickenpox()


```

Output:
```
0.07044873460147996
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```